---
rule_files:
  - /tmp/rules.verify

group_eval_order:
  - recordingRules.rules
  - alerts.rules

tests:
# CnaoDown positive tests
  - interval: 1m
    input_series:
      - series: "kubevirt_cnao_operator_up"
        values: "0 0 0 0 0 0"

    alert_rule_test:
      - eval_time: 5m
        alertname: CnaoDown
        exp_alerts:
          - exp_annotations:
              summary: "CNAO pod is down."
              runbook_url: "https://kubevirt.io/monitoring/runbooks/CnaoDown"
            exp_labels:
              severity: "warning"
              operator_health_impact: "warning"
              kubernetes_operator_part_of: "kubevirt"
              kubernetes_operator_component: "cluster-network-addons-operator"
# CnaoDown negative tests
  - interval: 1m
    input_series:
      - series: "kubevirt_cnao_operator_up"
        values: "1 1 1 0 1 1"
      - series: "kubevirt_cnao_operator_up"
        values: "1 1 1 1 1 1"
      - series: "kubevirt_cnao_operator_up"
        values: "0 0 0 0 0 1"

    alert_rule_test:
      - eval_time: 5m
        alertname: CnaoDown
        exp_alerts:

# NetworkAddonsConfigNotReady positive tests
  - interval: 1m
    input_series:
      - series: "kubevirt_cnao_cr_ready{namespace='ci'}"
        values: "0 0 0 0 0 0"

    alert_rule_test:
      - eval_time: 5m
        alertname: NetworkAddonsConfigNotReady
        exp_alerts:
          - exp_annotations:
              summary: "CNAO CR NetworkAddonsConfig is not ready."
              runbook_url: "https://kubevirt.io/monitoring/runbooks/NetworkAddonsConfigNotReady"
            exp_labels:
              severity: "warning"
              operator_health_impact: "warning"
              kubernetes_operator_part_of: "kubevirt"
              kubernetes_operator_component: "cluster-network-addons-operator"

# NetworkAddonsConfigNotReady negative tests
  - interval: 1m
    input_series:
      - series: "kubevirt_cnao_cr_ready{namespace='ci'}"
        values: "1 1 1 0 1 1"
      - series: "kubevirt_cnao_cr_ready{namespace='ci'}"
        values: "1 1 1 1 1 1"
      - series: "kubevirt_cnao_cr_ready{namespace='ci'}"
        values: "0 0 0 0 0 1"
      - series: "kubevirt_cnao_cr_ready{namespace='not-ci'}"
        values: "0 0 0 0 0 0"

    alert_rule_test:
      - eval_time: 5m
        alertname: NetworkAddonsConfigNotReady
        exp_alerts:

# KubeMacPoolDuplicateMacsFound positive tests
  - interval: 1m
    input_series:
      - series: "kubevirt_kmp_duplicate_macs{namespace='ci'}"
        values: "1 1 1 1 1 1"

    alert_rule_test:
      - eval_time: 5m
        alertname: KubeMacPoolDuplicateMacsFound
        exp_alerts:
          - exp_annotations:
              summary: "Duplicate macs found."
              runbook_url: "https://kubevirt.io/monitoring/runbooks/KubeMacPoolDuplicateMacsFound"
            exp_labels:
              severity: "warning"
              operator_health_impact: "warning"
              kubernetes_operator_part_of: "kubevirt"
              kubernetes_operator_component: "cluster-network-addons-operator"

# KubeMacPoolDuplicateMacsFound negative tests
  - interval: 1m
    input_series:
      - series: "kubevirt_kmp_duplicate_macs{namespace='ci'}"
        values: "1 0 0 0 0 0"
      - series: "kubevirt_kmp_duplicate_macs{namespace='not-ci'}"
        values: "1 1 1 1 1 1"

    alert_rule_test:
      - eval_time: 5m
        alertname: KubeMacPoolDuplicateMacsFound
        exp_alerts:

# KubemacpoolDown positive tests
  - interval: 1m
    input_series:
      - series: "kubevirt_cnao_kubemacpool_manager_up"
        values: "0 0 0 0 0 0"
      - series: "kubevirt_cnao_cr_kubemacpool_aggregated"
        values: "1 1 1 1 1 1"

    alert_rule_test:
      - eval_time: 5m
        alertname: KubemacpoolDown
        exp_alerts:
          - exp_annotations:
              summary: "KubeMacpool is deployed by CNAO CR but KubeMacpool pod is down."
              runbook_url: "https://kubevirt.io/monitoring/runbooks/KubemacpoolDown"
            exp_labels:
              severity: "critical"
              operator_health_impact: "critical"
              kubernetes_operator_part_of: "kubevirt"
              kubernetes_operator_component: "cluster-network-addons-operator"

  # KubemacpoolDown negative tests
  - interval: 1m
    input_series:
      - series: "kubevirt_cnao_kubemacpool_manager_up"
        values: "0 0 0 0 1 0"
      - series: "kubevirt_cnao_cr_kubemacpool_aggregated"
        values: "1 1 1 1 1 1"

    alert_rule_test:
      - eval_time: 5m
        alertname: KubemacpoolDown
        exp_alerts:

  - interval: 1m
    input_series:
      - series: "kubevirt_cnao_kubemacpool_manager_up"
        values: "0 0 0 0 0 0"
      - series: "kubevirt_cnao_cr_kubemacpool_aggregated"
        values: "1 1 1 1 0 1"

    alert_rule_test:
      - eval_time: 5m
        alertname: KubemacpoolDown
        exp_alerts:

  - interval: 1m
    input_series:
      - series: "kubevirt_cnao_kubemacpool_manager_up"
        values: "1 1 1 1 1 1"
      - series: "kubevirt_cnao_cr_kubemacpool_aggregated"
        values: "0 0 0 0 0 0"

    alert_rule_test:
      - eval_time: 5m
        alertname: KubemacpoolDown
        exp_alerts:
